<a href="https://colab.research.google.com/github/Ayombi/dsc-strings/blob/master/index.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Run this cell without changes
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt

df = pd.read_csv("/heroes_information (1).csv", index_col=0)
df.replace("-", np.nan, inplace=True)

df
```

Output:
```
name  Gender Eye color               Race        Hair color  \
0             A-Bomb    Male    yellow              Human           No Hair   
1         Abe Sapien    Male      blue      Icthyo Sapien           No Hair   
2           Abin Sur    Male      blue            Ungaran           No Hair   
3        Abomination    Male     green  Human / Radiation           No Hair   
4            Abraxas    Male      blue      Cosmic Entity             Black   
..               ...     ...       ...                ...               ...   
729  Yellowjacket II  Female      blue              Human  Strawberry Blond   
730             Ymir    Male     white        Frost Giant           No Hair   
731             Yoda    Male     brown     Yoda's species             White   
732          Zatanna  Female      blue              Human             Black   
733             Zoom    Male       red                NaN             Brown   

     Height          Publisher Skin color Alignment  Weight  
0    
```

## 1. Filter Data to Relevant Columns

Modify `df` so that it only contains the relevant columns for our analysis. These columns are listed as a list of strings for your convenience below.

When you are finished with this question, `df` should have the same number of rows as before (734 rows) but only 3 columns.

***Hint:*** if you are seeing the correct output but the test is failing, make sure that you are actually modifying `df`, either by reassigning it (`df = `...) or by using `inplace=True` (where applicable).

# Pandas Data Cleaning Checkpoint

This checkpoint is designed to test your understanding of the content from the Pandas Data Cleaning Cumulative Lab.

Specifically, this will cover:

* Using pandas to filter data
* Using pandas to handle missing values
* Using matplotlib to create a graph using data from a dataframe

## Your Task: Analyze Superhero Eye Color

### Data Understanding

In this repository under the file path `./heroes_information.csv` there is a CSV file containing information about various characters from superhero media properties.

The features of interest for this analysis will be:

`name`: The name (or AKA) of the character

`Eye color`: The eye color of the character

`Alignment`: "good", "bad", or "neutral". For the purposes of this analysis, we will only consider those with "good" alignment to be "superheroes"

### Requirements

#### 1. Filter Data to Relevant Columns

#### 2. Filter Data to Relevant Rows

#### 3. Drop Rows with Missing Values

#### 4. Find the Top 5 Most Common Eye Colors

#### 5. Plot a Bar Chart of Eye Colors

### Setup

In the cell below we import the relevant libraries, open up the CSV file as a dataframe called `df`, and convert cells containing `-` into cells containing `NaN` (because `-` was used to represent missing data in the original dataset).

***Hint:*** If you ever accidentally drop data that you didn't mean to drop, you can come back to this cell and re-run it to load a fresh copy of the data. Before submission, make sure you restart the kernel and run all of the cells to make sure that everything works in order.

```python
df.shape
```

Output:
```
(734, 10)
```

```python
df.columns
```

Output:
```
Index(['name', 'Gender', 'Eye color', 'Race', 'Hair color', 'Height',
       'Publisher', 'Skin color', 'Alignment', 'Weight'],
      dtype='object')
```

```python
df.loc[:, ['name', 'Eye color', 'Alignment']]
```

Output:
```
name Eye color Alignment
0             A-Bomb    yellow      good
1         Abe Sapien      blue      good
2           Abin Sur      blue      good
3        Abomination     green       bad
4            Abraxas      blue       bad
..               ...       ...       ...
729  Yellowjacket II      blue      good
730             Ymir     white      good
731             Yoda     brown      good
732          Zatanna      blue      good
733             Zoom       red       bad

[734 rows x 3 columns]
```

```python
# Run this cell without changes
relevant_columns = [
    "name",
    "Eye color",
    "Alignment"
]

df[relevant_columns]
```

Output:
```
name Eye color Alignment
0             A-Bomb    yellow      good
1         Abe Sapien      blue      good
2           Abin Sur      blue      good
3        Abomination     green       bad
4            Abraxas      blue       bad
..               ...       ...       ...
729  Yellowjacket II      blue      good
730             Ymir     white      good
731             Yoda     brown      good
732          Zatanna      blue      good
733             Zoom       red       bad

[734 rows x 3 columns]
```

<!-- grade_id: cell-013bedf98205d5a5, points: 0, grade: false, solution: true -->
```python
df[relevant_columns]



df
```

Output:
```
name  Gender Eye color               Race        Hair color  \
0             A-Bomb    Male    yellow              Human           No Hair   
1         Abe Sapien    Male      blue      Icthyo Sapien           No Hair   
2           Abin Sur    Male      blue            Ungaran           No Hair   
3        Abomination    Male     green  Human / Radiation           No Hair   
4            Abraxas    Male      blue      Cosmic Entity             Black   
..               ...     ...       ...                ...               ...   
729  Yellowjacket II  Female      blue              Human  Strawberry Blond   
730             Ymir    Male     white        Frost Giant           No Hair   
731             Yoda    Male     brown     Yoda's species             White   
732          Zatanna  Female      blue              Human             Black   
733             Zoom    Male       red                NaN             Brown   

     Height          Publisher Skin color Alignment  Weight  
0    
```

<!-- grade_id: cell-12f91738928bc996, points: 1, grade: true, solution: false -->
```python
# df should now have 734 rows and 3 columns
assert df.shape == (734, 3)

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-14-276ae96bec96> in <cell line: 2>()
      1 # df should now have 734 rows and 3 columns
----> 2 assert df.shape == (734, 3)
      3 
      4 # PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
      5 # THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

AssertionError:
```

## 2. Filter Data to Relevant Rows

Now, modify `df` so that it only contains rows where `Alignment` is `"good"`.

When you are finished with this question, `df` should still have 3 columns, but fewer rows.

```python

```

<!-- grade_id: cell-a85288413b87a128, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df = df[df['Alignment']=='good']

df
```

Output:
```
name  Gender Eye color            Race        Hair color  \
0             A-Bomb    Male    yellow           Human           No Hair   
1         Abe Sapien    Male      blue   Icthyo Sapien           No Hair   
2           Abin Sur    Male      blue         Ungaran           No Hair   
6        Adam Monroe    Male      blue             NaN             Blond   
7       Adam Strange    Male      blue           Human             Blond   
..               ...     ...       ...             ...               ...   
728     Yellowjacket    Male      blue           Human             Blond   
729  Yellowjacket II  Female      blue           Human  Strawberry Blond   
730             Ymir    Male     white     Frost Giant           No Hair   
731             Yoda    Male     brown  Yoda's species             White   
732          Zatanna  Female      blue           Human             Black   

     Height          Publisher Skin color Alignment  Weight  
0     203.0      Marvel Comics        NaN
```

<!-- grade_id: cell-524e5947d92b492e, points: 1, grade: true, solution: false -->
```python
# df should now have 496 rows and 3 columns
assert df.shape == (496, 3)

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-16-11727544014a> in <cell line: 2>()
      1 # df should now have 496 rows and 3 columns
----> 2 assert df.shape == (496, 3)
      3 
      4 # PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
      5 # THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

AssertionError:
```

## 3. Drop Rows with Missing Values

Now that all rows in the dataset are superheroes (alignment of "good"), note that we are missing the eye color for some of them:

```python
# Run this cell without changes
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 496 entries, 0 to 732
Data columns (total 10 columns):
 #   Column      Non-Null Count  Dtype  
---  ------      --------------  -----  
 0   name        496 non-null    object 
 1   Gender      477 non-null    object 
 2   Eye color   378 non-null    object 
 3   Race        279 non-null    object 
 4   Hair color  382 non-null    object 
 5   Height      496 non-null    float64
 6   Publisher   484 non-null    object 
 7   Skin color  35 non-null     object 
 8   Alignment   496 non-null    object 
 9   Weight      495 non-null    float64
dtypes: float64(2), object(8)
memory usage: 42.6+ KB
```

```python

```

Modify `df` so that it only contains rows where `Eye color` is not missing (i.e. is not NaN).

Again, when you are finished with this question, `df` should still have 3 columns, but fewer rows again.

<!-- grade_id: cell-f100fd02624d2c85, points: 0, grade: false, solution: true -->
```python
df.dropna()

df
```

Output:
```
name  Gender Eye color               Race Hair color  \
1            Abe Sapien    Male      blue      Icthyo Sapien    No Hair   
2              Abin Sur    Male      blue            Ungaran    No Hair   
39            Archangel    Male      blue             Mutant      Blond   
41               Ardina  Female     white              Alien     Orange   
74                Beast    Male      blue             Mutant       Blue   
75            Beast Boy    Male     green              Human      Green   
114               Blink  Female     green             Mutant    Magenta   
149        Captain Atom    Male      blue  Human / Radiation     Silver   
166             Century    Male     white              Alien      White   
226              Domino  Female      blue              Human      Black   
227           Donatello    Male     green             Mutant    No Hair   
232        Dr Manhattan    Male     white     Human / Cosmic    No Hair   
233  Drax the Destroyer    Male       red   
```

<!-- grade_id: cell-a25ce1c12b55c2da, points: 1, grade: true, solution: false -->
```python
# df should now have 378 rows and 3 columns
assert df.shape == (378, 3)

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

## 4. Find the Top 5 Most Common Eye Colors

Create a variable `top_eye_colors` that represents the top 5 most common eye colors in the dataset. Then create a variable `top_eye_color_counts` that represents the corresponding frequencies.

`top_eye_colors` should be a list of strings, ordered from the most common to the 5th most common. `top_eye_color_counts` should be a list of integers, ordered from the highest count to the 5th highest.

***Hint:*** If you have a list-like data structure (e.g. a pandas series or NumPy array), you can easily convert it to a Python list using `list()`.

<!-- grade_id: cell-a98fb450d5f9d47d, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code

top_eye_colors = None
top_eye_color_counts = None

# YOUR CODE HERE
raise NotImplementedError()

print("Colors:", top_eye_colors)
print("Counts:", top_eye_color_counts)
```

<!-- grade_id: cell-89189e9369383ce5, points: 0.5, grade: true, solution: false -->
```python
# top_eye_colors should be a list of strings
assert type(top_eye_colors) == list
assert type(top_eye_colors[0]) == str

# There should be 5 eye colors in top_eye_colors
assert len(top_eye_colors) == 5

# "blue" should be one of the eye colors in the list
assert "blue" in top_eye_colors

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

<!-- grade_id: cell-86273f7d0b77b084, points: 0.5, grade: true, solution: false -->
```python
# top_eye_color_counts should be a list of integers
assert type(top_eye_color_counts) == list
assert type(top_eye_color_counts[0]) == int or type(top_eye_color_counts[0]) == np.int64

# There should be 5 values in top_eye_color_counts
assert len(top_eye_color_counts) == 5

# 167 should be one of the counts in the list
assert 167 in top_eye_color_counts

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

## 5. Plot a Bar Chart of Popular Eye Colors

Create a matplotlib figure called `fig` containing a labeled bar chart with the number of superheroes who have each of the top 5 most popular eye colors (as encoded in `top_eye_colors` and `top_eye_color_counts`). You can find bar chart documentation [here](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.bar.html#matplotlib.axes.Axes.bar).

### Alternative Data

If you were not able to complete the previous question, you can use these hard-coded eye colors:

```
red
orange
yellow
green
blue
```

And these hard-coded counts:

```
46
2
19
73
167
```

***Note:*** there is no need to use the above values if you were able to complete the previous question and you have valid `top_eye_colors` and `top_eye_color_counts` variables.

### Starter Code

The starter code creates a figure called `fig` and axes called `ax`. Use those variables in your solution in order to pass the test.

**DO NOT** use the `df.plot` interface to answer this question. Use the `fig` and `ax` variables provided.

Set the *axes* title to `bar_chart_title` specified below.

<!-- grade_id: cell-143d320c2f71e081, points: 0, grade: false, solution: true -->
```python
bar_chart_title = "Top 5 Most Common Superhero Eye Colors"

fig, ax = plt.subplots()

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-700ecd741b83691e, points: 1, grade: true, solution: false -->
```python
# The axis should contain 5 bars
assert len(ax.containers[0]) == 5

# One of the x tick labels should be "blue"
tick_text = [tick.get_text() for tick in ax.get_xticklabels()]
assert "blue" in tick_text

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```