<a href="https://colab.research.google.com/github/Ayombi/dsc-strings/blob/master/Phase_1_Code_Challenge_Notebook.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Phase 1 Code Challenge
This code challenge is designed to test your understanding of the Phase 1 material. It covers:

- Pandas
- Data Visualization
- Exploring Statistical Data
- Python Data Structures

*Read the instructions carefully.* Your code will need to meet detailed specifications to pass automated tests.

## Code Tests

We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. To determine what the issue is, read the comments in the code test cells, the error message you receive, and the item instructions.

---
## Part 1: Pandas [Suggested Time: 15 minutes]
---
In this part, you will preprocess a dataset from the video game [FIFA19](https://www.kaggle.com/karangadiya/fifa19), which contains data from the players' real-life careers.

```python
# Run this cell

import pandas as pd
import numpy as np
from numbers import Number
import warnings
warnings.filterwarnings('ignore')
```

### 1.1) Read `fifa.csv` into a pandas DataFrame named `df`

Use pandas to create a new DataFrame, called `df`, containing the data from the dataset in the file `fifa.csv` in the folder containing this notebook.

Hint: Use the string `'./fifa.csv'` as the file reference.

**Starter Code**

    df = pd.read_csv('/content/fifa (2).csv')

<!-- grade_id: cell-cf9f259ed8f4b909, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df = pd.read_csv('/content/fifa (2).csv')
```

<!-- grade_id: cell-40c6c91273593f1f, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a DataFrame named df

assert type(df) == pd.DataFrame

```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 18207 entries, 0 to 18206
Data columns (total 88 columns):
 #   Column                    Non-Null Count  Dtype  
---  ------                    --------------  -----  
 0   ID                        18207 non-null  int64  
 1   Name                      18207 non-null  object 
 2   Age                       18207 non-null  int64  
 3   Photo                     18207 non-null  object 
 4   Nationality               18207 non-null  object 
 5   Flag                      18207 non-null  object 
 6   Overall                   18207 non-null  int64  
 7   Potential                 18207 non-null  int64  
 8   Club                      17966 non-null  object 
 9   Club Logo                 18207 non-null  object 
 10  Value                     18207 non-null  object 
 11  Wage                      18207 non-null  object 
 12  Special                   18207 non-null  int64  
 13  Preferred Foot            18159 non-null  object 
 14  Intern
```

### 1.2) Convert the `'Release Clause'` values from Euros to dollars

The `'Release Clause'` variable contains prices denominated in Euros. Use the exchange rate `1 Euro = 1.2 Dollars` to convert the prices to dollars.

<!-- grade_id: cell-7dea27dfb0840749, points: 0, grade: false, solution: true -->
```python
# Define the exchange rate
euro_to_dollar_exchange_rate = 1.2
# Convert the 'Release Clause' values from Euros to dollars
df['Release Clause'] = df['Release Clause'] * euro_to_dollar_exchange_rate
df['Release Clause'].head()
```

Output:
```
0    326160.0
1    183024.0
2    328464.0
3    199584.0
4    282816.0
Name: Release Clause, dtype: float64
```

<!-- grade_id: cell-722ee65f73aaf88b, points: 1, grade: true, solution: false -->
```python
# Ignore this cell

```

### 1.3) Drop rows from `df` with missing values for the `'Release Clause'` feature.
    
Make sure that `df` remains the name of the dataset with the dropped rows.

<!-- grade_id: cell-389eaeac18624fad, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df = df.dropna(subset=['Release Clause'])
```

<!-- grade_id: cell-d9fc548551ee3f7e, points: 1, grade: true, solution: false -->
```python
# This test confirms that your dataset has the correct number of observations after dropping

assert df['Release Clause'].shape[0] == 16643

```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 16643 entries, 0 to 18206
Data columns (total 88 columns):
 #   Column                    Non-Null Count  Dtype  
---  ------                    --------------  -----  
 0   ID                        16643 non-null  int64  
 1   Name                      16643 non-null  object 
 2   Age                       16643 non-null  int64  
 3   Photo                     16643 non-null  object 
 4   Nationality               16643 non-null  object 
 5   Flag                      16643 non-null  object 
 6   Overall                   16643 non-null  int64  
 7   Potential                 16643 non-null  int64  
 8   Club                      16643 non-null  object 
 9   Club Logo                 16643 non-null  object 
 10  Value                     16643 non-null  object 
 11  Wage                      16643 non-null  object 
 12  Special                   16643 non-null  int64  
 13  Preferred Foot            16643 non-null  object 
 14  Intern
```

### 1.4) Create a list `top_10_countries` containing the names of the 10 countries with the most players (using the `'Nationality'` column).

Hint: Your answer should include England, Germany, Spain, France, and Argentina

**Starter Code**

    top_10_countries =

<!-- grade_id: cell-f2a7bdcc98878531, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
top_10_countries = df['Nationality'].value_counts().head(10).index.tolist()
print(top_10_countries)
```

Output:
```
['England', 'Germany', 'Spain', 'France', 'Argentina', 'Brazil', 'Italy', 'Colombia', 'Japan', 'Netherlands']
['England', 'Germany', 'Spain', 'France', 'Argentina', 'Brazil', 'Italy', 'Colombia', 'Japan', 'Netherlands']
```

<!-- grade_id: cell-a5bfa46eb38629b9, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a list named top_10_countries

assert type(top_10_countries) == list

# This test confirms that top_10_countries contains England, Germany, Spain, France, and Argentina

assert set(['England', 'Germany', 'Spain', 'France', 'Argentina']).issubset(set(top_10_countries))

```

## Part 2: Data Visualization [Suggested Time: 20 minutes]
This part uses the same FIFA dataset, and asks you to plot data using `matplotlib`.

```python
# Run this cell

import matplotlib
import matplotlib.pyplot as plt

%matplotlib inline
```

### 2.1) Create a matplotlib figure `player_count_figure` containing a labeled bar chart with the number of players from England, Germany, Spain, France, and Argentina

Use the strings provided below (`bar_chart_title`, `bar_chart_count_label`, and `bar_chart_series_label`) to title and label your bar chart.

Hint: These are the countries with the top 5 numbers of players, so you may be able to adapt some of the code you used for question 1.4. If you were unable to complete 1.4, use the following values:

```
Country Name  | Num Players
============  | ===========
England       | 1000
Germany       | 900
Spain         | 800
France        | 700
Argentina     | 600
```

**Starter Code**

    player_count_figure, ax = plt.subplots(figsize=(10, 6))

<!-- grade_id: cell-248dd5aab2bb2f87, points: 0, grade: false, solution: true -->
```python
bar_chart_countries = ['England', 'Germany', 'Spain', 'France', 'Argentina']

bar_chart_title = '5 Countries with the Most Players'
bar_chart_count_label = 'Number of Players'
bar_chart_series_label = 'Nationality'

# YOUR CODE HERE
player_count_figure, ax = plt.subplots(figsize=(10, 6))
ax.plot(top_10_countries, Num Players)
ax.legend(["Sample Data"]);
```

Output:
```
File "<ipython-input-22-09ca7ee1d46f>", line 9
    ax.plot(top_10_countries, Num Players)
                              ^
SyntaxError: invalid syntax. Perhaps you forgot a comma?
```

<!-- grade_id: cell-128f21e6161d6c70, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a figure named player_count_figure

assert type(player_count_figure) == plt.Figure

# This test confirms that the figure contains exactly one axis

assert len(player_count_figure.axes) == 1

```

<!-- grade_id: cell-dec8d410badba377, points: 1, grade: true, solution: false -->
```python
# These tests confirm that the figure has a title and axis labels

assert player_count_figure.axes[0].get_title() != ''
assert player_count_figure.axes[0].get_ylabel() != ''
assert player_count_figure.axes[0].get_xlabel() != ''

```

### 2.2) Create a matplotlib figure `tackle_figure` containing a labeled scatter plot visualizing the relationship between `StandingTackle` (on X axis) and `SlidingTackle` (on Y axis)

Use the strings provided below (`scatter_plot_title`, `standing_tackle_label`, and `sliding_tackle_label`) to title and label your scatter plot.

**Starter Code**

    tackle_figure, ax = plt.subplots(figsize=(10, 6))

<!-- grade_id: cell-064cad4a79f53237, points: 0, grade: false, solution: true -->
```python
scatter_plot_title = 'Relationship Between Standing Tackles and Sliding Tackles'
standing_tackle_label = 'Standing Tackles'
sliding_tackle_label = 'Sliding Tackles'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7fc1fc24025435b8, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a figure named tackle_figure

assert type(tackle_figure) == plt.Figure

# This test confirms that the figure contains exactly one axis

assert len(tackle_figure.axes) == 1

```

## Part 3: Exploring Statistical Data [Suggested Time: 20 minutes]
This part does some exploratory analysis using the same FIFA dataset.

### 3.1) Create numeric variables `mean_age` and `median_age` containing the mean and median player ages (respectively).

**Starter Code**

    mean_age =
    median_age =

<!-- grade_id: cell-236fe375f6c9327c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f71047bc869336cf, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created numeric variables named mean_age and median_age

assert isinstance(mean_age, Number)
assert isinstance(median_age, Number)

```

### 3.2) Create numeric variables `oldest_argentine_name` and `oldest_argentine_age` containing the name and age (respectively) of the oldest player with Argentina nationality.

**Starter Code**

    oldest_argentine_name =
    oldest_argentine_age =

<!-- grade_id: cell-0c07eea0bdc0dfa4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-469be575f68081da, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a string variable named oldest_argentine_name

assert type(oldest_argentine_name) == str

# This test confirms that you have created a numeric variable named oldest_argentine_age

assert isinstance(oldest_argentine_age, Number)

```

## Part 4: Python Data Structures [Suggested Time: 20 min]

Below is a dictionary `players` with information about soccer players. The keys are player names and the values are dictionaries containing each player's age, nationality, and a list of teams they have played for.

```python
# Run this cell

players = {
    'L. Messi': {
        'age': 31,
        'nationality': 'Argentina',
        'teams': ['Barcelona']
    },
    'Cristiano Ronaldo': {
        'age': 33,
        'nationality': 'Portugal',
        'teams': ['Juventus', 'Real Madrid', 'Manchester United']
    },
    'Neymar Jr': {
        'age': 26,
        'nationality': 'Brazil',
        'teams': ['Santos', 'Barcelona', 'Paris Saint-German']
    },
    'De Gea': {
        'age': 27,
        'nationality': 'Spain',
        'teams': ['Atletico Madrid', 'Manchester United']
    },
    'K. De Bruyne': {
        'age': 27,
        'nationality': 'Belgium',
        'teams': ['Chelsea', 'Manchester City']
    }
}
```

### 4.1) Create a list `player_names` of all the player names in dictionary `players`.

**Starter Code**

    player_names =

<!-- grade_id: cell-7c9289df331d0247, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5d33a8eee819cb46, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a list named player_names

assert type(player_names) == list

```

### 4.2) Create a list of tuples `player_nationalities` containing each player's name along with their nationality.

**Starter Code**

    player_nationalities =

<!-- grade_id: cell-f14953191eda7561, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e0f87d9c0c4cbb1, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a list named player_nationalities

assert type(player_nationalities) == list

```

### 4.3) Define a function `get_players_on_team()` that returns a list of names of all the players who have played on a given team.

Your function should take two arguments:

- A dictionary of player information
- A string containing a team name (for which you are trying to find the player names)

**Starter Code**

    def get_players_on_team(player_dict, team_name):
        player_list = []
    
        return player_list

<!-- grade_id: cell-365f0240380ce928, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e74522c0d595f7fd, points: 2, grade: true, solution: false -->
```python
# This test confirms that get_players_on_team() returns the right names for Manchester United

manchester_united_players = ['Cristiano Ronaldo', 'De Gea']
players_on_manchester_united = get_players_on_team(players, 'Manchester United')

assert players_on_manchester_united == manchester_united_players

```